Megakaryocytes assemble a three-dimensional cage of extracellular matrix that controls their maturation and anchoring to the vascular niche

Author response:

The following is the authors’ response to the previous reviews.

Minor Issues:

(1) As the authors mention, MKs have been suggested to mature rapidly at the sinusoids, and both integrin KO and laminin KO MKs appear mislocalized away from the sinusoids. Additionally, average MK distances from the sinusoid may also help separate whether the maturation defects could be in part due to impaired migration towards CXCL12 at the sinusoid. Presumably, MKs could appear mislocalized away from the sinusoid given the data presented suggesting they are leaving the BM and entering circulation. Additional commentary on intrinsic (ex-vivo) MK maturation phenotypes may help strengthen the author's conclusions

Thank you for your insightful suggestion regarding intrinsic MK maturation defects in integrin KO and laminin KO mice. This indeed could be the case. We have now addressed this possibility in the revised discussion section (page 14; lines 14-15), acknowledging intrinsic maturation defects as a potential contributor to observed maturation issues.

(2) It would be helpful if the authors could comment as to whether MKs are detectable in blood.

We appreciate the opportunity to clarify this point. Intact Itgb1-/-/Itgb3-/- MKs were not detected in the peripheral blood by either flow cytometry or blood smear analysis. This indicates that megakaryocytes do not normally circulate in the systemic bloodstream. Instead, we observed large MK nuclei trapped specifically within the lung capillaries, consistent with their known physical retention in the pulmonary circulation during platelet release. This explanation is now better explained on page 10, lines 14-19.

(3) Supplementary Figure 6 - shows no effect on in vitro MK maturation and proplt, or MK area - But Figures 6B/6C demonstrate an increase in total MK number in MMP-inhibitor treated mice compared to control. This discrepancy should be better discussed.

We have now expanded the discussion in the revised manuscript to address the different results obtained in vitro and in vivo, emphazing that the in vitro model may not fully recapitulate the complex and dynamic bone marrow ECM niche. Additionally, differences in the source and regulation of MMPs likely contribute to the differing outcomes, underlining the importance of studying these processes within their physiological context. For instance, non-megakaryocytic sources of MMPs and paracrine regulatory mechanisms may play a critical role within the physiological microenvironment, ultimately affecting MK proliferation and maturation in a manner not observed in simplified culture systems. This clarifications can be found on page 12, lines 6-17.

(4) A function of the ECM discussed relates to MK maturation but in the B1/3 integrin KO mice, the presence of the ECM cage is reduced but there appears to be no significant impact upon maturation (Supplementary Figure 4). By contrast, MMP inhibition in vivo (but not in vitro) reduces MK maturation. These data could be better clarified in the text.

Thank you for raising this important point. While Suppl. Figure 4 shows normal size and ploidy in DKO MK, a critical defect is revealed at the ultrastructural level. Mature DKO MKs exhibit severe dysplasia of the demarcation membrane system (DMS), characterized by extensive membrane accumulation and abnormal archirecture, with no typical platelet territories visible. This DMS defect directly impairs MK maturation and explains the thrombocytopenia observed in these mice. Increased emperipolesis further indicated disrupted maturation processes. These observations confirm the essential role of the ECM cage in supporting proper DMS organization and overall MK maturation in vivo, consistent with findings from MMP inhibition experiments. We have clarified and emphasized the significance of these DMS abnormalities in the revised manuscripts, including updated results (Page 9, lines 17-21) and a new EM image in Suppl. Figure 4.

Reviewer #1 (Public review):

The authors report on a thorough investigation of the interaction of megakaryocytes (MK) with their associated ECM during maturation. They report convincing evidence to support the existence of a dense cage-like pericellular structure containing laminin γ1 and α4 and collagen IV, which interacts with integrins β1 and β3 on MK and serve to fix the perisinusoidal localization of MK and prevent their premature intravasation. As with everything in nature, the authors support a Goldilocks range of MK-ECM interactions - inability to digest the ECM via inhibition of MMPs leads to insufficient MK maturation and development of smaller MK. This important work sheds light into the role of cell-matrix interactions in MK maturation, and suggests that higher-dimensional analyses are necessary to capture the full scope of cellular biology in the context of their microenvironment. The authors have responded appropriately to the majority of my previous comments.

We sincerely thank the reviewer for their insightful comments.

Some remaining points:

In a previous critique, I had suggested that "it is unclear how activation of integrins allows the MK to become "architects for their ECM microenvironment" as the authors posit. A transcriptomic analysis of control and DKO MKs may help elucidate these effects". The authors pointed out the technical difficulty of obtained sufficient numbers of MK for such analysis, which I accept, and instead analyzed mature platelets, finding no difference between control and DKO platelets. This is not necessarily surprising, since mature circulating platelets have no need to engage an ECM microenvironment, and for the same reason I would suggest that mature platelet analyses are not representative of MK behavior as regards ECM interactions.

We fully agree with the reviewer that platelet analyses do not accurately reflect the behavior of MKs in the context of interactions with the ECM. This understanding is also one of the reasons why we chose not to include RT-PCR data on platelets in our manuscript. Instead, we emphasize the role of integrins as essential regulators of ECM remodeling, as they transmit traction forces that can significantly influence this process. We also report reduced RhoA activation in DKO MK, which is likely to affect ECM organization. We believe that these explanations contribute to a clearer understanding of how integrin activation enables megakaryocytes to act as "architects" of their ECM microenvironment.

Reviewer #2 (Public review):

This study makes a significant contribution to understanding the microenvironment of megakaryocytes (MKs) in the bone marrow, identifying an extracellular matrix (ECM) cage structure that influences MK localization and maturation. The authors provide compelling evidence for the presence of this ECM cage and its role in MK homeostasis, employing an array of sophisticated imaging techniques and molecular analyses.The authors have addressed most of the concerns raised in the previous review, providing clarifications and additional data that strengthen their conclusion.

More broadly, this work adds to a growing recognition of the ECM as an active participant in haematopoietic cell regulation in the bone marrow microenvironment. This work could pave the way to future studies investigating how the megakaryocytes' ECM cage affects their function as part of the haematopoietic stem cell niche, and by extension, influences global haematopoiesis.

We thank this reviewer for providing such constructive feedback.

Changes from the earlier version to the revised version

Abstract

Abstract

Megakaryocytes, the progenitor cells of blood platelets, play a crucial role in hemostasis by residing in the bone marrow and ensuring continuous platelet production. Unlike other hematopoietic cells, megakaryocytes do not enter the blood circulation intact. They remain anchored within the bone marrow while extending cytoplasmic protrusions called proplatelets through the sinusoidal endothelial barrier. These proplatelets subsequently fragment into functional platelets. This unique process of intravasation facilitates efficient platelet production while maintaining the megakaryocyte cell body within the bone marrow niche, thus preventing potential thrombotic complications. How the extracellular matrix (ECM) influences the delicate balance between megakaryocyte retention and proplatelet extension remains largely unknown. Here, we investigate the spatial organization and functional role of ECM components in the megakaryocyte vascular niche. Our findings reveal that laminin and collagen IV form three-dimensional (3D) ECM cages encompassing megakaryocytes and anchor them to the sinusoidal basement membrane. Gene deletion shows the existence of laminin α4 in the ECM cage that is necessary to maintain megakaryocyte-sinusoid interactions. Notably, megakaryocytes actively contribute to the ECM cage assembly; β1/β3 integrin knockout weakens these structures, increasing intravasation and entire megakaryocyte entry into circulation. The retention of megakaryocytes by these 3D ECM cages depends on dynamic remodeling processes. Inhibition of ECM proteolysis results in denser cage formation, increasing the frequence of immature megakaryocytes with impaired demarcation membrane system (DMS) development. Thus, the ECM cage represents a novel concept of an active and dynamic 3D microenvironment that is continuously remodeled and essential for maintaining megakaryocyte perivascular positioning. This specific microarchitecture guides megakaryocyte maturation and intravasation, underscoring the critical role of ECM microarchitecture and dynamics in megakaryocyte function.

Introduction

Megakaryocytes, the precursors of blood platelets, differentiate from hematopoietic stem cells and mature in a complex bone marrow microenvironment. Megakaryocytes are large, highly polyploid cells that contain a complex invaginated network of membranes, known as the demarcation membrane system (DMS), which serves as a membrane reservoir for platelet production. In contrast to other hematopoietic cells, megakaryocytes exhibit a unique platelet production process. They do not enter the bloodstream intact. Instead, they remain anchored in the bone marrow and extend cytoplasmic protrusions called proplatelets through the sinusoidal endothelial barrier. These protrusions protrude into the lumen of bone marrow sinusoids, where they undergo fragmentation to release individual platelets into the circulation. This distinctive mechanism allows for the efficient production and release of platelets while maintaining the megakaryocyte cell body within the bone marrow microenvironment, thus reducing the risk of thrombotic complications (Boscher et al., 2020; Stone et al., 2022). The remaining cell body, composed of a nucleus surrounded by a thin rim of cytoplasm, is ultimately phagocytosed in the stroma (Radley & Haller, 1983).

Megakaryocytes are strategically located at the interface between the bone marrow and the blood circulation, specifically at the parasinusoidal region (Lichtman et al., 1978; Stegner et al., 2017). Intravasation, the coordinated passage of megakaryocyte fragments through the sinusoidal barrier, requires an original, complex, and dynamic adaptation of megakaryocytes to highly different microenvironments, both mechanically and biologically. They are positioned next to the endothelial lining and its underlying basement membrane, in equilibrium between the constrained 3D environment of the bone marrow and the fluid environment of the blood. In this context, they are in contact with two types of ECM organization : basement membrane and interstitial ECM.

Megakaryocytes actively influence their structural microenvironment through several ECM remodeling mechanisms, including the endocytosis of plasma proteins like fibrinogen, and the synthesis and release of extracellular matrix (ECM) components such as fibronectin laminin, and type IV collagen (Abbonante et al., 2017; Handagama et al., 1987; Malara et al., 2014). Additionally, they produce ECM-modifying proteins, including matrix metalloproteinases (MMPs), tissue inhibitors of metalloproteinases (TIMPs), and lysyl oxidase (LOX), which facilitate the continuous renewal and regulation of the matrix (Eliades et al., 2011; Malara et al., 2018; Villeneuve et al., 2009). Furthermore, we have recently discovered that megakaryocytes could remodel the substrate-bound fibronectin matrix into basal fibrillar structures surrounding the cells. The extent of fibrillogenesis depended on the stiffness of the substrate, and relied on both β1 and β3 integrins (Guinard et al., 2023). This intricate interplay between megakaryocytes and their surrounding ECM remodeling demonstrates the dynamic nature of the vascular niche and highlights its critical role in platelet production.

The ECM surrounding megakaryocytes plays a crucial role in their development and function. Numerous studies have identified essential components of this matrix, including collagen IV, fibronectin, fibrinogen, and laminin (Larson & Watson, 2006; Malara et al., 2014; Semeniak et al., 2016; Susek et al., 2018). The presence and function of collagen I fibers remain a subject of ongoing debate, with some researchers proposing that they guide megakaryocyte proplatelets toward sinusoids (Oprescu et al., 2022). In vitro studies have shown that while collagen I may inhibit proplatelet formation, other ECM components, like collagen IV, fibrinogen, and fibronectin, can facilitate this process (Balduini et al., 2008; Malara et al., 2014; Sabri et al., 2004). The diverse and sometimes contradictory roles of these ECM components in megakaryocyte behavior underscore the complexity of this research area and highlight the need for further investigation, particularly in vivo.

The present study provides new insight into the spatial organization of the ECM that envelops megakaryocytes, elucidating its specific architecture and molecular mechanisms governing its dynamics and function. We performed advanced 2D and 3D analyses of mouse bone marrow and identified a novel ECM cage containing laminin chains ϒ1 and α4, and collagen IV that anchored megakaryocytes to the sinusoidal basement membrane. Deleting laminin α4 impairs the connections between megakaryocytes and sinusoids. Megakaryocytes act actively in forming and maintaining this ECM cage, relying on their integrins β1 and β3. Indeed, deleting these integrins causes significant impairment of the cage’s integrity, resulting in increased intravasation of megakaryocytes and an unexpected release of intact megakaryocytes into the bloodstream. Moreover, in vivo, matrix metalloproteinases (MMPs) inhibition reveals that dynamic ECM remodeling is crucial for maintaining the cage structure and supporting megakaryocyte development. In conclusion, our study unveils a 3D ECM cage that physically stabilizes megakaryocytes within their vascular niche. This structure facilitates the physiological regulation of megakaryocyte maturation and intravasation at the bone marrow-bloodstream interface.

Results

Laminin and collagen IV create a 3D cage around sinusoid-associated megakaryocytes

The organization of the ECM around megakaryocytes in the vascular niche is not well understood due to challenges in in vivo observation. To investigate this, we utilized immunofluorescence (IF) microscopy on ultrathin bone marrow cryosections (250 nm thickness). This approach offers superior axial resolution compared to traditional confocal microscopy, enabling high-resolution localization of ECM components. Megakaryocytes were visualized using antibodies against GPIbβ, while sinusoids were identified with antibodies against FABP4 (Figure 1A). Our findings revealed that laminin ϒ1 chains and collagen IV delineated the basement membrane underlying the endothelium of sinusoids and the outer contour of mature megakaryocytes (Figure 1A, Suppl. Figure 1Aa). Adjacent cells exhibited granular staining for laminin ϒ1 and collagen IV, which may indicate a potential contribution to the ECM constitution (arrow in Suppl. Figure 1Ag-h). Fibronectin was detected in the basement membrane and around megakaryocytes, while fibrinogen was more widespread, associated with the basement membrane, the megakaryocyte surface, and intracellular granules (Suppl. Figure 1Ab-c). Von Willebrand factor (VWF) signal was restricted to the alpha granules of megakaryocytes (Suppl. Figure 1Ad). We observed no detectable signals for type I and III collagen around megakaryocytes or sinusoids despite using antibodies validated on positive controls (Suppl. Figure 1Ae-f, Suppl. Figure 1B-C).

Lamininϒ1 and collagen IV form 3D cages of ECM for megakaryocytes, directly connected to the sinusoidal basement membrane

(A) Left panel. Schematic representation of the experimental workflow for 2D imaging of immunostained bone marrow cryosections from WT mice. Confocal imaging is performed on single ultrathin sections with an axial resolution of 250 µm. Right panel. Representative 2D images of a sinusoid-associated megakaryocyte immunostained for laminin ϒ1 (red), GPIbβ (white) and FABP4 (cyan). Cell nuclei are visualized with DAPI (blue) (from one out of three independent IF experiments).

(B) Left panel. Schematic representation of the experimental workflow for 3D analysis of whole-mount bone marrow preparations from WT mice. A stack of confocal images covering half the depth of the megakaryocyte is acquired and then z-projected to create a maximum projection image. Right panel. Representative maximal projection images of sinusoid-associated megakaryocyte immunostained for laminin ϒ1 (red), GPIbβ (white) and FABP4 (cyan) (from one out of three independent IF experiments).

(C) Representative maximal projection images of sinusoid-associated megakaryocyte immunostained for collagen IV (green) and GPIbβ (white). The inset image shows the FABP4 (cyan) and megakaryocyte (white) immunostaining.

(D) Bone marrow-isolated megakaryocyte maintaining an ECM cage. Left panel. Schematic of the experimental procedure used to isolate mouse bone marrow megakaryocytes. Right panel. Maximal projection 3D images showing the persistence of the ECM cage (collagen IV in green) around a freshly isolated megakaryocyte (GPIbβ in white).

*, sinusoid lumen ; arrowheads, basement membrane-cage connections; BM, bone marrow; bm, basement membrane ; MK, megakaryocyte; Bars, 10 µm.

Having established the expression of laminin ϒ1, collagen IV, fibronectin, and fibrinogen in the direct megakaryocyte microenvironment, their spatial organization was next investigated using 3D imaging of whole-mount bone marrow preparations (Figure 1B). Analysis of maximum Z-stack projections revealed that laminin ϒ1 and collagen IV formed a reticular 3D ECM cage completely enveloping megakaryocytes and extending radially from the sinusoidal basement membrane (Figure 1B-C). With this approach, the basement membrane-ECM cage connections were clearly visible when examining the full z-stack series. Examples are shown in movies 1 and 2 and in Suppl. Figures 1D-E. Quantification showed that almost all megakaryocytes near sinusoids (sMK, 92.8 ± 3.3 %) have a laminin cage, compared to only 11.4 ± 4.8 % of megakaryocytes in the parenchyma (pMK) (Suppl. Figure 1F-G). Parenchymal megakaryocytes (pMK), which represented only 18 ± 1.3 % of all megakaryocyte population (Suppl. Figure 1H), were instead surrounded by a a sparse thin network of laminin ϒ1 (Suppl. Figure 1F). Althought we cannot exclude that ECM cage can be form on its own, our data suggests that ECM cage assembly may require interactions between megakaryocytes and the sinusoidal basement membrane. Parenchymal megakaryocytes (pMK), which represented only 18 ± 1.3 % of all megakaryocyte population (Suppl. Figure 1H), were instead surrounded by a a sparse thin network of laminin ϒ1 (Suppl. Figure 1F). The ECM cage was present at all stages of megakaryocyte maturation, including megakaryocytes with proplatelet extension (Suppl. Figures 1I-J). Remarkably, after mechanical dissociation and size exclusion, nearly half of the megakaryocytes successfully retained their cages (53.4 ± 5.6 %, 329 megakaryocytes counted from 3 mice), indicating strong physical attachments between both components (Figure 1D). While fibronectin and fibrinogen are present around megakaryocytes and at the vessel-cell interface, they do not form a reticular ECM cage. Furthermore, no connection was found between fibronectin and fibrinogen deposition with the sinusoid basement membrane, in contrast to the findings for laminin and collagen IV (Suppl. Figure 1K). These observations demonstrate that megakaryocytes establish precise and tight interactions with an ECM cage made of laminin ϒ1 and collagen IV in a spatially confined microenvironment at the sinusoidal basement membrane.

Reduced 3D laminin cage and vessel-associated megakaryocytes in Lamα4-/- mouse bone marrow

Among the γ1 chain-bearing laminin isoforms, laminin α4 was abundantly found in bone marrow sinusoid basement membrane (Susek et al., 2018). To explore the direct influence of ECM organization on platelets and megakaryocytes, we utilized laminin α4-deficient mice (Lamα4-/-). These mice exhibit mild thrombocytopenia, with platelet counts approximately 20% lower than wild-type (WT) mice (Cai et al., 2022). Their platelets maintain a normal ultrastructure and functions, as shown by flow cytometry and EM analysis (Suppl. Figures 2A-C). Bone marrow from Lamα4-/- mice showed a normal number of mature megakaryocytes with characteristic ultrastructural features (Figure 2C and Suppl. Figure 2D). Analysis of bone marrow explants further demonstrated that these megakaryocytes could extend proplatelets, indicating preserved fundamental function (Suppl. Figures 2 E-F).

Reduced laminin ϒ1 cage and megakaryocyte-sinusoid interactions in Lamα4-/- mouse bone marrow.

(A-B) Depletion of laminin α4 leads to a reduction in the laminin ϒ1 deposition, but not in collagen IV, in the cage around megakaryocytes and in the sinusoid basement membrane.

A. Representative maximal projection images showing the immunostaining of laminin γ1 (red) or collagen IV (green) in the Lamα4-/- compared to control mice. Two magnifications are shown for Lama4-/- mouse.

B. Quantification of laminin γ1 and collagen IV surface coverage per megakaryocyte and per basement membrane surface (laminin: 6<cage<17 and 15<bm<17; collagen IV: 6<cage<16 and 11<bm<21 expressed as a percentage, ****P>0.001 unpaired t-Test).

(C-E) Depletion of laminin α4 leads to a decrease in the sinusoid-associated megakaryocytes.

C. Representative maximal projections showing the immunostaining of laminin γ1 (red) and megakaryocytes (GPIbβ in white) in the bone marrow of Lamα4-/- and control mice.

D. Quantification of the total number of megakaryocytes per surface unit (s.u., 12,945 μm2, n=3 for each genotype, 87<n<92, P=0.9228, ns, Mann-Whitney test).

E. Quantification of the sinusoid-associated megakaryocytes (n=3, ****P>0.001 unpaired t-Test) in control (grey) and Lamα4-/- (dark) mice. Arrows point to megakaryocytes which are not associated with sinusoids (pMK).

bm, basement membrane; MK, megakaryocyte; pMK, MK in the parenchyma; Bar, 10 µm.

We next examined the ECM cage using an antibody against laminin γ1. In Lamα4-/- mice, laminin γ1 deposition around megakaryocytes and in the sinusoid basement membrane decreased by 1.7- and 2.6-fold compared to controls, indicating a disruption of the laminin cage (Figures 2A-B, upper panels). However, collagen IV organization remained unaltered around the megakaryocytes in these mice (Figure 2A-B, lower panels), showing that laminin α4 is not necessary for collagen IV cage formation. We also analyzed megakaryocyte localization in the bone marrow of Lamα4-/- mice (Figure 2C). The total megakaryocyte density was similar to that of control mice (Figure 2D), but we observed more parenchymal MK with a 1.5-fold reduction in the number of megakaryocytes near sinusoids (Figure 2E). These observations highlights the role of the laminin α4 cage in maintaining optimal positioning of megakaryocytes near sinusoids.

Integrins maintain the structural properties of the ECM cages

Integrins play a crucial role in ECM remodeling. Megakaryocytes express β1 and β3 integrins as main ECM receptors (Yang et al., 2022). To elucidate the molecular mechanisms governing the intricate interactions between megakaryocytes and the ECM cage, we investigate the subcellular localization of integrins on immunostained ultrathin bone marrow cryosections. Using conformation-independent antibodies, we showed the presence of β1 subunit at on the plasma membrane and on the DMS (Figure 3A). Similar findings were obtained for total β3 integrins (Suppl. Figure 3A). In contrast, the activated form of β1 integrin was expressed only at interaction sites with laminin γ1 and collagen IV, suggesting that activated β1 integrins mediate megakaryocyte interactions with the ECM cage (Figure 3B). Activated β3 integrin involvement remains to be determined, as no specific signal was detected when testing two different antibodies (JonA-PE, Pac 1).

Integrin-mediated control of the 3D ECM cage around megakaryocytes

(A) Representative 2D images of Pf4cre bone marrow cryosections (250 nm) showing a sinusoid-associated megakaryocyte immunostained for β1 integrin (MAB1997 in yellow). Right: the boxed area is shown at a higher magnification.

(B) Activated β1 integrins form functional adhesion structures around megakaryocyte surfaces. Representative 2D images of Pf4cre bone marrow cryosections showing laminin (red, upper panels) or collagen IV (green, lower panels) staining around a sinusoid-associated megakaryocyte immunostained for activated β1 integrin (9EG7 in cyan) and GPIbβ (white). Right: Magnification of the boxed area showing co-localization of the ECM proteins and activated β1 integrin.

(C) Depletion of β1 and β3 integrins leads to a reduction in the laminin deposition on the surface of megakaryocytes. Representative 3D images showing a decrease in laminin deposition (red) on Itgb1-/-/Itgb3-/- megakaryocytes compared to Pf4cre.

(D) Quantification of laminin surface coverage per megakaryocyte (in %, 17<n<19 as indicated in the bars, ***P>0.001 unpaired t-Test), (Middle) expression profile of the laminin staining along straight-lines (25 µm long) visible as white lines in the confocal images, and (Right) quantification of mesh sizes (in µm, 14<n<16, **P<0.01 Mann-Whitney).

*, sinusoid lumen; MFI, mean fluorescence intensity; MK, megakaryocyte; pm, plasma membrane; Bars, 10 µm.

To test if integrin-mediated signaling plays a role in the structural assembly of the ECM cage, a transgenic mouse model lacking β1 and β3 integrins in the megakaryocyte lineage was used (Itgb1-/-/Itgb3-/-) (Figure 3C). Itgb1-/-/Itgb3-/- mice were generated crossing the two single loxP-flanked lines (β1fl/fl and β3fl/fl mice) with Pf4-Cre mice (expressing the Cre recombinase under the control of the Pf4 promoter). Our analysis revealed a 2.6-fold reduction in laminin γ1 deposition on the megakaryocyte surface in Itgb1-/-/Itgb3-/- mice compared to Pf4-Cre control mice. Furthermore, intensity profiles demonstrated that laminin γ1 forms a less dense network around the megakaryocytes in the double knockout mice, with significantly increased mesh sizes (10.6 ± 1.3 µm vs 6.3 ± 0.6 µm, respectively) (Figure 3D). Notably, this lower ECM density surrounding megakaryocytes in Itgb1-/-/Itgb3-/- mice is not linked to increased matrix degradation or reduced laminin synthesis (Suppl. Figures 3B-C). As expected, the laminin γ1 network at the basement membrane remained unaffected in these mice, as the integrin deletion is restricted to megakaryocytes (Suppl. Figures 3E-F).

Further confocal and immunoEM examination revealed a decrease in the expression of fibrillar fibronectin around Itgb1-/-/Itgb3-/-megakaryocytes, along with mislocalization of fibrinogen in the extracellular space of the DMS instead of the granules (Suppl. Figures 3G and 3H). These findings show that megakaryocytes, through β1 and β3 integrins dependant mechanisms, participate in their ECM microenvironment remodelling (ECM cage and fibronectin, fibrinogen deposition) and that megakaryocytes β1 and β3 integrins are required.

Normal ECM organization is essential for proper megakaryocyte intravasation

Earlier version: Next, we investigate the relevance of the ECM cage to platelet and megakaryocyte functions. Flow cytometry showed that Itgb1-/-/Itgb3-/-platelets have intact α-granule secretion but impaired integrin-mediated aggregation, consistent with the bleeding phenotype observed in these mice (Janus-Bell et al., 2024) (Suppl. Figure 4A). As we previously described that Itgb1-/-/Itgb3-/- mice had a 50% reduction in platelet count (Guinard et al., 2023), we hypothesized that the organization of the ECM cage could contribute to the maturation process of megakaryocytes. Firstly, the number, ploidy, and maturation stages of megakaryocytes in Itgb1-/-/Itgb3-/-mice are similar to control mice. However, consistent with previous reports, these mutant megakaryocytes have a disorganized membrane system and smaller cytoplasmic areas compared to normal cells (Suppl. Figures 4B-E). We also analyzed megakaryocyte localization in the bone marrow tissue. We found a significantly higher proportion of megakaryocytes that were extending intravascular fragments (i.e., intravasation) in Itgb1-/-/Itgb3-/- mice (15.1 ± 2.7% of total bone marrow megakaryocytes) compared to that in Pf4Cre mice (2.9 ± 0.4% of total megakaryocytes). More strikingly, intact megakaryocytes were found in the sinusoid lumen, a rare phenomenon in control mice under physiological conditions (Figures 4A-B).

Revised version: Next, we investigate the relevance of the ECM cage to platelet and megakaryocyte functions. Flow cytometry showed that Itgb1-/-/Itgb3-/-platelets have intact α-granule secretion but impaired integrin-mediated aggregation, consistent with the bleeding phenotype observed in these mice (Janus-Bell et al., 2024) (Suppl. Figure 4A). As we previously described that Itgb1-/-/Itgb3-/- mice had a 50% reduction in platelet count (Guinard et al., 2023), we hypothesized that the organization of the ECM cage could contribute to the maturation process of megakaryocytes. Firstly, the number, size and ploidy megakaryocytes in Itgb1-/-/Itgb3-/- mice are similar to control mice. However, these mutant megakaryocytes show increased emperipolesis and severe DMS dysplasia, characterized by excessive membrane accumulation and abnormal archirecture. Typical platelet territories are poorly defined, indicating disrupted maturation processes (Suppl. Figures 4B-E). We also analyzed megakaryocyte localization in the bone marrow tissue. We found a significantly higher proportion of megakaryocytes that were extending intravascular fragments (i.e., intravasation) in Itgb1-/-/Itgb3-/- mice (15.1 ± 2.7% of total bone marrow megakaryocytes) compared to that in Pf4Cre mice (2.9 ± 0.4% of total megakaryocytes). More strikingly, intact megakaryocytes were found in the sinusoid lumen, a rare phenomenon in control mice under physiological conditions (Figures 4A-B).

Integrins protects megakaryocytes from entering the bloodstream as whole cells

(A-C)) Higher proportion of intravasation events in mice lacking β1/β3 integrins.

A. Representative confocal images of Pf4cre and Itgb1-/-/Itgb3-/-whole-mount bone marrow immunostained for GPIbβ (white) and FABP4 (cyan).

B. Quantification of megakaryocyte intravasation and circulating megakaryocytes (3 mice minimum for each genotype, 213<n<397 for Pf4cre and Itgb1-/-/Itgb3-/-, ** P<0.001 Mann Whitney).

C. Quantification of the laminin γ1 deposition in the ECM cage in single knockout integrins and in GPVI knockout. Quantification of the intravasation events in single knock-out mice showing that both integrins are essential for the proper anchoring of megakaryocytes in their vascular niche.

(D -E) Intravital two-photon imaging of Itgb1-/-/Itgb3-/- mouse calvarial bone marrow

D. Tissues were stained with intravenously injected AF488-conjugated anti-GPIX antibody and rhodamin dextran. The white arrow indicates an intrasinusoidal Itgb1-/-/Itgb3-/-megakaryocyte, dotted lines illustrate the sinusoid wall and the values in the left corner show the time-lapses.

E. Quantification of circulating megakaryocytes, expressed as a percentage of the total number of megakaryocyte (from 3 independent experiments, 130<n<136, 0.0279 *P<0.1, Paired t-test).

(F-G) Large megakaryocyte nuclei detected in the pulmonary capillaries of Itgb1-/-/Itgb3-/- mice

F. Representative confocal images showing megakaryocytes’ nucleus (arrow, GPIbβ green, DAPI nucleus) within the pulmonary microvessels of Itgb1-/-/Itgb3-/- mice. Cyan dotted lines indicate the vessel wall.

G. Quantification of the intravascular megakaryocytes (from 5 independent experiments, 28<n<149, 0.0079 **P<0.01, Mann-Whitney).

(H) Two TEM image showing intravascular entire Itgb1-/-/Itgb3-/- megakaryocyte.

*, sinusoid wall; FG, fibrinogen; FN, fibronectin; FG, fibrinogen; n, nucleus; sMK, sinusoid-associated MK; pMK, MK in the parenchyma; PPT, proplatelets; Bars in A-G, 10 μm; Bar in E, 5 µm; Bar in G, 30µm.

To elucidate which of the two integrins was responsible for the observed phenotype, we employed single-knockout mice (Figure 4C). Analysis of the ECM cage revealed a trend towards decreased laminin ϒ1 density around Itgb1-/- megakaryocytes, while no notable changes were observed around Itgb3-/-cells. No statistically significant increase in intravasation events was observed in either single knock-out mouse model. This suggests a redundant role of both integrins in regulating the megakaryocyte ECM cage and the intravasation process, potentially compensating for each other’s role. Additionally, we investigate the role of GPVI, one of the two well-characterized collagen receptors in this process (Figure 4C), and no alterations in ECM cage formation or intravasation behavior were noticed, which is in agreement with previous reports (Semeniak et al., 2019).

Earlier version: To understand the unusual localization of circulating Itgb1-/-/Itgb3-/-megakaryocytes in the bloodstream, we used intravital 2-photon microscopy imaging to describe the dynamic of the megakaryocyte behavior. We used the GPIX marker, a subunit of the GPIb-V-IX complex that is expressed on mature megakaryocytes, to track the behavior of megakaryocyte in living mice. Among the stabilized megakaryocytes at the parasinusoidal interface, we could observe the cellular distortions of Itgb1-/-/Itgb3-/- megakaryocytes and their exit of the marrow (Figures 4D-E, movie 3). Flow cytometry analysis did not allow us to detect intact megakaryocytes in the blood in Itgb1-/-/Itgb3-/-blood samples (Suppl. Figure 4F). However, large megakaryocyte nuclei were observed within the downstream pulmonary capillaries of these mice (Figures 4F-G), suggesting that circulating megakaryocytes in Itgb1-/-/Itgb3-/-mice are in large proportion retained in lung microvasculature shortly after entering the bloodstream. Transmission electron microscopy (TEM) observation confirmed that intact Itgb1-/-/Itgb3-/- megakaryocytes were similar in size and ultrastructure to those in the stroma compartment (Figure 4H). Furthermore, no significant change in the size of the endothelial pores (Itgb1-/-/Itgb3-/-: 4.6 ± 0.3 µm; Pf4Cre: 4.3 ± 0.4 µm) was observed, indicating that increased megakacaryocyte intravasion in not link to endothelium alteration in mutant mice.

Revised version: To understand the unusual localization of circulating Itgb1-/-/Itgb3-/-megakaryocytes in the bloodstream, we used intravital 2-photon microscopy imaging to describe the dynamic of the megakaryocyte behavior. We used the GPIX marker, a subunit of the GPIb-V-IX complex that is expressed on mature megakaryocytes, to track the behavior of megakaryocyte in living mice. Among the stabilized megakaryocytes at the parasinusoidal interface, we could observe the cellular distortions of Itgb1-/-/Itgb3-/- megakaryocytes and their exit of the marrow (Figures 4D-E, movie 3). Intact Itgb1-/-/Itgb3-/-megakaryocytes were not detected in the peripheral blood by either flow cytometry or blood smear analysis (Suppl. Figure 4F), indicating that megakaryocytes do not normally circulate in the systemic bloodstream. Instead, large MK nuclei were found trapped specifically in the lung capillaries (Figures 4F–G), consistent with their known physical retention in the pulmonary circulation where they can release platelet. Transmission electron microscopy (TEM) observation confirmed that intact Itgb1-/-/Itgb3-/-megakaryocytes were similar in size and ultrastructure to those in the stroma compartment (Figure 4H). Furthermore, no significant change in the size of the endothelial pores (Itgb1-/-/Itgb3-/-: 4.6 ± 0.3 µm; Pf4Cre: 4.3 ± 0.4 µm) was observed, indicating that increased megakacaryocyte intravasion in not link to endothelium alteration in mutant mice.

Integrins promote megakaryocyte adhesion to the ECM components of the bone marrow

We next assessed the contribution of integrins adhesive function in megakaryocyte anchorage to ECM. Under static conditions, most Itgb1-/-/Itgb3-/-megakaryocytes did not spread and still showed a round shape on immobilized laminin, fibronectin or fibrinogen (Figures 5A-B). We next measured the adhesive potential of freshly isolated bone marrow megakaryocytes to test for such an adhesive role. To this end, we used a miniaturized microfluidic-based experimental model in which individual megakaryocyte detachment was tracked when exposed to a flow rate of 300s-1, similar to the flow typically found in sinusoids (Suppl. Figure 5A-B). Pf4-Cre megakaryocytes had capture yields of 70.5% on laminin, 83.6% on fibrillar fibronectin, and 82.0% on fibrinogen. Laminin exhibited the least adhesion, emphasizing the importance of the molecular composition of the said 3D ECM cage and its evident impact on the adhesion response of megakaryocytes in vivo (Figures 5C-D). Pf4-Cre megakaryocytes remained anchored, while Itgb1⁻/⁻/Itgb3⁻/⁻ megakaryocytes experienced higher detachment rates across various ECMs (movies 4 and 5). In line with this conclusion, we used the bone marrow explant model to study the adhesion properties of the megakaryocytes associated with their 3D ECM cage (Figures 5E-F). This model enabled us to quantify the physical detachment of megakaryocytes at the periphery of the explants. After 3 hours, more megakaryocytes detached from Itgb1-/-/Itgb3-/-bone marrow tissue than Pf4-Cre explants, reaching a plateau at 6 h. These findings indicate that β1 and β3 integrins control ECM-megakaryocyte interactions in their native bone marrow microenvironment.

Integrins promote megakaryocyte adhesion to the ECM components of the bone marrow

(A-B) Impaired adhesion and spreading of Itgb1-/-/Itgb3-/- megakaryocytes.

A. Representative SEM images depicting bone marrow-derived megakaryocytes adhering on laminin.

B. Spreading (hatched bars) and round (grey bars) megakaryocytes were counted following 3h incubation on laminin, fibronectin (FN), and fibrinogen (FG) (in %) (from 4-6 independent experiments).

(C-D) Microfluidic flow chamber to study megakaryocyte adhesion efficiency.

C. Representatives bright field images showing that upon flow application, Itgb1-/-/Itgb3-/- megakaryocytes detach from fibrillary fibronectin protein, while Pf4Cre MKs remain attached.

D. Quantification of the detachment of Pf4cre and Itgb1-/-/Itgb3-/- megakaryocytes on laminin, fibrillar fibronectin and fibrinogen (from 5 to 7 independent experiments,**P<0.01, ***P<0,001, Mann-Whitney).

(E-F) Reduced physical anchoring of Itgb1-/-/Itgb3-/- megakaryocytes to bone marrow.

E. Representatives bright field images of the ten femur bone marrow sections placed in an incubation chamber (left panel), of the box (center panel) and of the megakaryocytes released from the periphery of the explants (right panel).

F. Quantification of the number of Pf4cre and Itgb1-/-/Itgb3-/- megakaryocytes released from the explants following 3h (from 10 to 13 independent experiments, 594<n<1095 for Pf4cre and Itgb1-/-/Itgb3-/-, *P<0.05, unpaired t-test).

dotted lines, MK detachment; MK, megakaryocytes, FN, fibronectin; FG, fibrinogen; n, number of cells studied; Bars in A, 10 μm; Bars in B, 20 µm; Bars in B, 30 µm.

Collectively, our results highlight the essential roles of β1 and β3 integrins in forming 3D ECM cages around megakaryocytes and modulating their adhesion within the bone marrow, which helps stabilize the cells in their vascular niche and prevent the passage of intact megakaryocytes through the sinusoid barriers.

Cage compression via metalloproteinase inhibition affects the maturation of megakaryocytes

Our results suggest that a weakened ECM cage promotes megakaryocyte intravasation through reduced adhesion. We, therefore, tested whether an increase in cage density also may affect megakaryocyte functions. For that purpose, we proposed to decrease the catabolic aspect occurring in the ECM equilibrium between its synthesis and degradation. To this aim, mice were injected with batimastat (30 mg/kg) and ilomastat (20 mg/kg) for 7 days, in order to inhibit the activation of MMPs in the megakaryocyte microenvironment (Winer et al., 2018) (Supplemental Figure 6A). These mice produced platelets normally in terms of both count and average platelet function (Suppl. Figures 6B-C). Regarding the cage, this treatment increased the ECM density in megakaryocyte surrounding, as evidenced by reduced fiber length and pore size (Figures 6A-B).

Maturation of the 3D ECM cage is correlated with maturation of the DMS in megakaryocytes

(A-B) MMP inhibition leads to a densification of the 3D ECM cage.

A. Representative 3D confocal images showing a significant increase in collagen IV deposition (green) on the megakaryocyte surface in treated mice treated with the intravenous cocktail of protease inhibitors (B + I).

B. (Left) Quantification of collagen IV fluorescence showed a shortening of collagen IV fibers in treated mice compared to that in control mice (from 3-5 independent experiments, 20<n<22 as indicated in the bars, ****P<0;001, Mann-Whitney), (Middle) Histograms of fluorescence intensity versus distance showed an increase in cross-linking with a reduction in pore size (white lines of 25µm length are visible in the confocal images), (Right) Reduction in mesh sizes in treated mice (from 3 independent experiments, 7<n<12, ***P>0.001, t-test).

(C-D) MMP inhibition affects megakaryocyte growth.

C. Representative confocal images from DMSO vs B+I treated mice immunostained for GPIbβ (white) and FABP4 (cyan), zoomed-in images showing the difference in megakaryocyte size between the two groups (arrows).

D. Quantification of the number of megakaryocytes per bone marrow area (194 x 194 µm) (from 3 independent experiments, 229<n<483 as indicated in the bars, ****P<0.0001, Mann-Whitney).

(E-F) MMP inhibition leads to an increased in the number of immature megakaryocytes.

E. TEM observation revealed the presence of numerous immature megakaryocytes (stage II) in treated mice as compared to fully mature megakaryocytes in control mice (stage III).

F. Quantification of the total number of megakaryocytes (145<n<169 as indicated in bars, **P<0.01, Mann-Whitney) and in the proportion of immature megakaryocytes (stage II) in the B + I group (from 3 independent experiments, **P < 0.05, ***P < 0.01, one way ANOVA with Tukey correction).

(G-H) MMP inhibition reduced release of megakaryocytes from the bone marrow explant.

G. Representatives bright field images of the megakaryocytes (arrows) released from the periphery of the control and treated explants.

H. Quantification of the number of megakaryocytes released following 3h and 6h (from 6 independent experiments, 369<n<783 for DMSO and B+I, *P<0.05, unpaired t-test).

B + I, batimastat + Ilomastat; sMK, sinusoid-associated MK; pMK, MK in the parenchyma; n, number of cells studied; Bars in A and C, 10 μm; Bars in E, 5 µm; Bar in G, 50 µm.

Earlier version: We investigated whether this denser ECM cage might affect megakaryopoiesis. Remarkably, treatment with MMP inhibitors resulted in a significantly higher proportion of smaller megakaryocytes compared to untreated mice (white arrows in Figures 6C-D). To better evaluate the extent of this defect, we employed TEM (Figure 6E). We found that 53.7% of the megakaryocytes lacked the appropriate organization of the DMS in MMP inhibitors-treated mice, reflecting a primary failure in the cytoplasmic maturation (Stage II in Figure 6F right). To determine whether MMP inhibition could directly affect megakaryocytes, we cultured primary megakaryocytes in the presence of increasing concentrations of inhibitors. We found no significant change in their proliferation, maturation, and proplatelet formation (Suppl. Figure 6D). Although these in vitro system and exogenous MMPs may not fully recapitulate the complexity of the bone marrow situation, these findings suggest that disturbing ECM remodelling significantly impacts megakaryocyte maturation.

Revised version: We investigated whether this denser ECM cage might affect megakaryopoiesis. Remarkably, treatment with MMP inhibitors resulted in a significantly higher proportion of smaller megakaryocytes compared to untreated mice (white arrows in Figures 6C-D). To better evaluate the extent of this defect, we employed TEM (Figure 6E). We found that 53.7 % of the megakaryocytes lacked the appropriate organization of the DMS in MMP inhibitors-treated mice, reflecting a primary failure in the cytoplasmic maturation (Stage II in Figure 6F right). These mice also showed a significant increase in the total megakaryocyte number compared to the control condition (Figures 6D and 6F left), consistent with altered bone marrow ECM remodeling affecting megakaryocyte retention. To assess whether MMP inhibition directly affects megakaryocyte maturation and density, we cultured primary megakaryocytes in the presence of increasing concentrations of inhibitors. We observed no significant change in proliferation, maturation, or proplatelet formation (Suppl. Figure 6D). While this in vitro model and the application of exogenous MMP inhibitors offer valuable mechanistic insights, they do not fully recapitulate the complexity and regulation of the bone marrow ECM in vivo. For example, MMPs from other cell types and paracrine signals in the bone marrow may influence megakaryocyte behavior. Nonetheless, these findings are consistent with the conclusion that disrupting ECM remodeling significantly impacts megakaryocyte maturation within the physiological bone marrow.

Earlier version: An additional phenotype associated with a denser ECM environment is a significant increase in the total megakaryocyte number in MMP inhibitor-treated mice compared to the control condition (Figures 6D and 6F left). Moreover, immunofluorescence analysis of bone marrow sections revealed an upregulation of activated β1 integrin (Suppl. Figure 6E). These results indicate that increased ECM density may promote megakaryocyte adhesion within the congested microenvironment of the bone marrow. To further investigate this point, we quantified the number of megakaryocytes released from bone marrow explants. We observed a significant reduction in megakaryocyte egress from the bone marrow microenvironment after 6 h in MMP inhibitor-treated compared to control mice (Figure 6G-H). Together, these results indicate that MMPs are essential regulators of ECM homeostasis within the bone marrow, influencing megakaryocyte maturation and attachment to the vascular niche.

Revised version: Immunofluorescence analysis of bone marrow sections revealed an upregulation of activated β1 integrin (Suppl. Figure 6E). These results indicate that increased ECM density may promote megakaryocyte adhesion within the congested microenvironment of the bone marrow. To further investigate this point, we quantified the number of megakaryocytes released from bone marrow explants. We observed a significant reduction in megakaryocyte egress from the bone marrow microenvironment after 6 h in MMP inhibitor-treated compared to control mice (Figure 6G-H). Together, these results indicate that MMPs are essential regulators of ECM homeostasis within the bone marrow, influencing megakaryocyte maturation and attachment to the vascular niche.

Discussion

Earlier version: Physiological platelet formation relies on the strategic positioning of mature megakaryocytes at sinusoids to facilitate polarised proplatelet extension and platelet release into circulation. This study identifies a 3D ECM cage anchoring megakaryocytes to bone marrow sinusoids. This cage, composed of laminin γ1 and α4 and collagen IV fibers, extends from the basement membrane and surrounds megakaryocytes. Megakaryocyte integrin signaling and ECM proteolysis regulate its remodeling and adhesive properties. This dynamic microenvironment is key for megakaryocyte positioning, maturation, and intravasation into the bloodstream, representing a novel concept in understanding physiological platelet formation (Graphical Abstract, Figure 7).

Revised version: Physiological platelet formation relies on the strategic positioning of mature megakaryocytes at sinusoids to facilitate polarised proplatelet extension and platelet release into circulation. This study identifies a 3D ECM cage anchoring megakaryocytes to bone marrow sinusoids. This cage, composed of laminin γ1 and α4 and collagen IV fibers, extends from the basement membrane and surrounds megakaryocytes. Megakaryocyte integrin signaling and ECM proteolysis regulate its remodeling and adhesive properties. This dynamic microenvironment is key for megakaryocyte positioning, maturation, and intravasation, representing a novel concept in understanding physiological platelet formation (Graphical Abstract, Figure 7).

Integrin-mediated signaling and MMP proteolysis regulate the matrix remodeling and the adhesive properties of the 3D ECM cage, which control megakaryocyte maturation and intravasation at the bone marrow-blood interface.

In line with previous studies (Guinard et al., 2023; Larson & Watson, 2006; Malara et al., 2014; Semeniak et al., 2016), we identified the presence of collagen IV, laminin γ1, fibronectin, and fibrinogen surrounding sinusoid-associated megakaryocytes. Notably, when analyzed using immunofluorescence and transmission electron microscopy (TEM), collagen I fibers were absent under steady-state conditions. The discrepancy between our results and those reported in the literature may be attributed to differences in staining methodologies or variations in the physiological context of the bone marrow. Furthermore, our findings indicate that GPVI-deficient megakaryocytes do not display alterations in ECM cage formation, vessel association and intravasation behavior. These observations support studies showing GPVI and α2β1 integrins are not critical for megakaryocyte function, underscoring potential redundancy in ECM-receptor interactions (Semeniak et al., 2019). Besides, non-protein ECM molecules, including glycosaminoglycans, have been shown to play an essential role in supporting megakaryocyte function, including maturation (Petrey et al., 2016; Piszczatowski et al., 2022; Vögtle et al., 2019).

We identify laminin and collagen IV as the principal structural components of this ECM cage. Consistent with previous research (Abbonante et al., 2016; Cai et al., 2022; Susek et al., 2018), these ECM proteins were detected intracellularly within megakaryocytes and in stromal cells situated in their vicinity. This indicates that the ECM cage proteins may be derived from autocrine synthesis and their production by neighbouring cells. Interestingly, while fibronectin and fibrinogen were found around megakaryocytes, they do not form a ECM cage, suggesting distinct functional roles, potentially related to megakaryocyte expansion in bone marrow fibrosis (Malara et al., 2019; Matsuura et al., 2020).

Earlier version: Our findings in Lamα4-deficient mice reveal that the ECM cage is critical for positioning megakaryocytes near sinusoids. It is established that CXCL12 is a most potent physiological chemoattractant for megakaryocytes and that the parasinusoidal location of megakaryocytes in the bone marrow is mediated by the CXCL12/CXCR4 interaction (Hamada et al., 1998; Wang et al., 2019). Other chemoattractive factors also include fibroblast growth factor 4 (FGF4) (Avecilla et al., 2004) and VWF (Ouzegdouh et al., 2018). Besides its potential role as a reservoir of such growth factors, we show that the ECM cage is also a crucial element in maintaining megakaryocytes in the proximity of sinusoids, which is essential for their efficient maturation and optimal platelet production. We also find the persistent presence of this ECM cage throughout megakaryocyte maturation, supporting previous findings that megakaryocyte development occurs primarily at the sinusoid (Gaertner et al., 2024; Lichtman et al., 1978; Stegner et al., 2017).

Revised version: Our findings in Lamα4-deficient mice reveal that the ECM cage is critical for positioning megakaryocytes near sinusoids. It is established that CXCL12 is a most potent physiological chemoattractant for megakaryocytes and that the parasinusoidal location of megakaryocytes in the bone marrow is mediated by the CXCL12/CXCR4 interaction (Hamada et al., 1998; Wang et al., 2019). Other chemoattractive factors also include fibroblast growth factor 4 (FGF4) (Avecilla et al., 2004) and VWF (Ouzegdouh et al., 2018). Besides its potential role as a reservoir of such growth factors, we show that the ECM cage is also a crucial element in maintaining megakaryocytes in the proximity of sinusoids, which is essential for their efficient maturation and optimal platelet production. We also find the persistent presence of this ECM cage throughout megakaryocyte maturation, supporting previous findings that megakaryocyte development occurs primarily at the sinusoid (Gaertner et al., 2024; Lichtman et al., 1978; Stegner et al., 2017). Together, these findings underscore the essential structural and functional role of the ECM cage in ensuring megakaryocyte localization at the sinusoid niche, which is indispensable for their proper maturation. However, we also acknowledge that intrinsic cellular defects may contribute to maturation impairments observed in laminin and integrin deficient models.

Earlier version: Our study highlights the indispensable roles of β1 and β3 integrins in maintaining megakaryocyte stability within the vascular niche. We demonstrate that these integrins are crucial for two fundamental processes : (1) the building of the 3D ECM cages, and (2) the ECM-mediated adhesion mechanisms to the vascular niche. Integrins biochemical and mechanical signal transduction play key roles in ECM remodelling (Larsen et al., 2006) and can generate the forces that may cause these changes (Lemmon et al., 2009). For instance, RhoA inhibition in Itgb1⁻/⁻/Itgb3⁻/⁻ megakaryocytes (Guinard et al., 2023) is likely contributing to the altered ECM remodeling. Nevertheless, we still need to understand how in megakaryocyte β1 or β3 integrins participate in ECM remodelling and ECM cage dynamics. Our work further demonstrates that in these mice, intact megakaryocytes can enter the general circulation. Of note, no intravascular megakaryocytes were found in Lamα4-deficient mice, which have a normal collagen IV cage but a compromised laminin cage. This indicates that integrity of all components of the cage is not needed to keep megakaryocytes from entering circulation. Importantly, circulating megakaryocytes have also been observed in humans during emergency megakaryopoiesis, such as in cases of infectious diseases, inflammatory conditions, and following cardiopulmonary bypass procedures (Frydman et al., 2023; Gelon et al., 2022; Puhm et al., 2023; Rapkiewicz et al., 2020). These circulating megakaryocytes have been implicated in thrombotic complications, underscoring the importance of integrin signaling-dependent mechanisms regulating megakaryocyte retention and release.

Revised version: Our study highlights the indispensable roles of β1 and β3 integrins in maintaining megakaryocyte stability within the vascular niche. We demonstrate that these integrins are crucial for two fundamental processes : (1) the building of the 3D ECM cages, and (2) the ECM-mediated adhesion mechanisms to the vascular niche. Integrins biochemical and mechanical signal transduction play key roles in ECM remodelling (Larsen et al., 2006) and can generate the forces that may cause these changes (Lemmon et al., 2009). For instance, RhoA inhibition in Itgb1⁻/⁻/Itgb3⁻/⁻ megakaryocytes (Guinard et al., 2023) is likely contributing to the altered ECM remodeling. Nevertheless, we still need to understand how exactly β1 or β3 integrins participate in ECM remodelling and ECM cage dynamics. Our work further demonstrates that in these mice, intact megakaryocytes can cross the sinusoidal vessel wall. Of note, no intravascular megakaryocytes were found in Lamα4-deficient mice, which have a normal collagen IV cage but a compromised laminin cage. This indicates that integrity of all components of the cage is not needed to keep megakaryocytes from entering circulation. Importantly, circulating megakaryocytes have also been observed in humans lungs during emergency megakaryopoiesis, such as in cases of infectious diseases, inflammatory conditions, and following cardiopulmonary bypass procedures (Frydman et al., 2023; Gelon et al., 2022; Puhm et al., 2023; Rapkiewicz et al., 2020). These circulating megakaryocytes have been implicated in thrombotic complications, underscoring the importance of integrin signaling-dependent mechanisms regulating megakaryocyte retention and release.

The precise role of the ECM cage in platelet production remains incompletely understood. In the Lamα4⁻/⁻ mice, a collagen-rich ECM cage persists with normal fibronectin deposition, whereas the Itgb1⁻/⁻/Itgb3⁻/⁻ model displays a much more severe phenotype, characterized by the loss of both the laminin cage and collagen, along with the absence of fibrillar fibronectin. The preserved collagen and fibronectin in Lamα4⁻/⁻ mice may allow for residual activation of signaling pathways - potentially through integrins or alternative mechanisms - compared to the Itgb1⁻/⁻/Itgb3⁻/model, where these matrix components are absent.

MMPs profoundly influence megakaryocyte maturation and their association with the vascular niche by modulating the density of ECM cage architecture. This finding highlights the dynamic nature of the bone marrow microenvironment and its impact on megakaryopoiesis. Dysregulation of ECM production, often associated with bone marrow pathologies, can lead to its uncontrolled accumulation. In myelofibrosis, for instance, megakaryocytes secrete transforming growth factor-β1 (TGF-β1) which stimulates excessive ECM production by stromal cells (Abbonante et al., 2016). This pathological ECM accumulation leads to abnormally smaller megakaryocytes that exhibit reduced platelet production (Gianelli et al., 2023; Malara et al., 2018; Sarachakov et al., 2023). Likewise, we observed smaller and immature megakaryocytes retained in the constrained ECM microenvironment of the bone marrow, associated with an upregulation of activated β1 integrin in treated mice. Interestingly, previous research by K. Hoffmeister’s group has demonstrated that proper megakaryocyte localization at the sinusoids depends on β1 integrin glycosylation. Specifically, the absence of galactosylation in β4galt1-/- megakaryocytes leads to hyperactivity of β1 integrin, which negatively impacts the formation of the DMS (Giannini et al., 2020). Building on these insights, our study shows the coordinated interplay between integrin-mediated signaling and MMP proteolysis, working together to tightly regulate the density and cross-linking properties of the ECM cage and thereby control megakaryocyte maturation at the bone marrow-blood interface. Intriguingly, despite these effects on megakaryocyte maturation, we observed no significant changes in circulating platelet counts, and importantly, platelet function remained preserved. This supports the idea that differences in ECM composition can influence the signaling environment and megakaryocyte maturation, but do not fully abrogate platelet function.

The influence of ECM features - like fiber length and pore size - on megakaryocyte biology is complex and not yet fully understood. Longer ECM fibers may help cells adhere better (Barriga et al., 2018; Dolega et al., 2021). Larger pores could make it easier for megakaryocytes to mature and to extend proplatelets at the bone marrow-blood interface. Also, the mechanical properties of the ECM, particularly its stiffness, have emerged as critical environmental determinants of megakaryocyte development and function (Abbonante et al., 2017; Aguilar et al., 2016; Guinard et al., 2023; Leiva et al., 2018). Notably, megakaryocytes possess mechanosensing capabilities, primarily mediated by the β3 integrin subunit, enabling them to detect and respond to changes in substrate stiffness (Guinard et al., 2023). The assembly of extracellular matrix (ECM) fibronectin around megakaryopoiesis involves a more complex interplay of integrins, with both β1 and β3 integrins playing essential roles in fibrillogenesis (Abbonante et al., 2024; Guinard et al., 2023). This cooperative relationship appears to be conserved across various cell types (Attieh et al., 2017; De Mets et al., 2019; Kyumurkov et al., 2023). This may explain why the deletion of both integrins is required to affect the 3D ECM cage and megakaryocyte behavior significantly. We can not exclude that other ECM ligands of β3 integrins, such as fibronectin, may contribute to the organization of collagen IV and laminin, influencing the overall architecture of the vascular niche (Petito & Gresele, 2024; Yang et al., 2022; Zeng et al., 2018). Indeed, it is known that the fibronectin matrix favors the deposition of other ECM proteins, such as collagens IV and various other glycoproteins (Saunders & Schwarzbauer, 2019). This hypothesis underscores the need for further investigation into the spatio-temporal dynamics of ECM component assembly in the megakaryocyte vascular niche.

Overall, our findings reveal the supportive role of the ECM cage in platelet biogenesis. Integrin-mediated ECM interactions are clearly crucial, as demonstrated by the 50% reduction in platelet counts observed in integrin double knockout mice. In summary, the ECM cage acts as a finely tuned facilitator, optimising the efficiency and precision of platelet production by guiding megakaryocytes to the right place at the right time.

Earlier version: Material and methods

Revised version: Materials and methods

Animals

We used wild type (WT) mice (C57BL/6J from C. River, L’Arbesle, France), Itgb1-/-/Itgb3-/- double knock out mice and Pf4cre mice aged 10 to 15 weeks. Pf4cre mice expressed the Cre recombinase under the control of the Pf4 promoter. The β1fl/fl (Potocnik et al., 2000) and β3fl/fl (Morgan et al., 2010) mice were crossed with mice expressing the Cre recombinase under the control of the Pf4 promoter to obtain inactivation in the MK lineage. Itgb1-/-/Itgb3-/-double knock out (KO) mice were crosses of the two single KO lines, as described in Guinard et al., 2023. GPVI -/- mice and Lamα4-/-mice were from Nieswandt and Qian labs (Bender et al., 2011; Cai et al., 2022). All animal studies were approved by the French legislation for animal experimentation and in accordance with the guide for the care and use of laboratory animals as defined by the European laws (Animal Facility Agreement C-67-482-10).

Chemicals

Dimethyl sulfoxide (2438), bovine serum albumin (900.011), saponin (47036), Triton X-100 (T8787), fibrinogen and batimastat (SML0041) (Sigma-Aldrich, Rueil-Malmaison, France), Ilomastat (GM6001, HY-15768) (MedChemExpress, Clinisciences, France), Dulbecco’s modified Eagle’s medium (DMEM), penicillin, streptomycin, and glutamine (Invitrogen, Cergy-Pontoise, France), Gelatin-Oregon Green 488 conjugate (G-13186) (Life Technologies), Laminin 511 recombinant (Biolamina) and fibronectin from human plasma from (341635, Calbiochem) were used in this study.

Antibodies

Earlier version: See Tables 1 and 2 in “Data Supplements” for details.

Revised version: See Tables 1 and 2 for details.

GPIX and GPIbβ are components of the GPIb-IX complex, identifying mature megakaryocytes (Lepage et al., 2000). The choice of marker used to identify megakaryocytes in different experiments is primarily based on technical considerations. Intravital experiments have been standardised using AF488-conjugated anti-GPIX to consistently identify mature megakaryocytes. The rest of the manuscript uses GPIbβ (GP1bβ) due to its strong, specific, bright staining.

Extrusion and Preservation of Murine Bone Marrow

Bone marrow extrusion, without leaving any residual bone, was obtained by carefully flushing the femurs of mice with PBS using a 21-gauge needle attached to a 10 mL syringe (see photo in Figure 1D) (Scandola et al., 2021). To maximize the preservation of the bone marrow integrity, we used a double fixation procedure (4 % paraformaldehyde and 0.4% glutaraldehyde for 1 hour) immediately after bone marrow extrusion, followed by embedding it in 4 % low-melting-point agarose to preserve as much as possible their three-dimensional architecture.

Isolation of freshly isolated bone marrow-derived megakaryocytes

Bone marrow megakaryocytes were dissociated using a 21G needle and filtered through a 40 µm Millipore filter. The filter was then rinsed in DMEM + 1% fœtal bovine serum, and the isolated megakaryocytes were adjusted to 300 cells/ml in the same medium.

2D confocal microscopy on bone marrow cryosections

This study explores the interactions between megakaryocytes and their immediate ECM microenvironment. Ultrathin cryosections were used for their superior axial resolution, offering a 2- to 3-fold improvement over conventional confocal microscopy, which facilitates analysis of signal superposition. Cryosections from WT extruded bone marrow were used to study the ECM composition and distribution of sinusoid-associated megakaryocytes. Bone marrow was fixed in a mix of 2 % paraformaldehyde-0.2 % glutaraldehyde in 0.2 M sodium cacodylate buffer for 1 hour. The fixed samples were infiltrated with 2.3 M sucrose and frozen in liquid nitrogen. Ultrathin cryosections of 250 nm were obtained at −110°C with a LEICA Ultracut UCT cryo-ultramicrotome (Leica Microsystems). For immunofluorescence staining, cryosections were labeled with primary antibodies and conjugated-second antibodies of the appropriate species and DAPI, as reported in Table 1. They were examined under a confocal microscope (TCS SP8, Leica) using the 63x objective with a numerical zoom of 4 (pixel size: 0.09µm). The bone marrow specimens from 3 mice were examined under identical conditions, using constant exposure and the same irrelevant antibodies. No fluorescence was detected using isotype-specific control IgG.

3D confocal analysis on bone marrow vibratome sections

Whole-mount bone marrow preparations were used to investigate i) the localization of megakaryocytes in the bone marrow and the lung and ii) the spatial organization of ECM around sinusoid-associated megakaryocytes. Megakaryocytes were classified based on their maturation stage: stage I (presence of granules, no clear DMS visible), stage II (developing DMS not yet organized), and stage III (DMS organized in platelet territories). Mice from each genotype were analyzed and image acquisitions were performed in a blinded manner.

For megakaryocyte localization, sections of femurs’s bone marrow or lungs fixed in 2% paraformaldehyde and 0.2% glutaraldehyde for 1 hour, with a thickness of 250 µm, were incubated with FABP4 overnight. This was followed by overnight staining with Alexa 568-conjugated anti-GPIbβ antibody and Alexa 488 Donkey anti-Goat for FABP4. Finally, Hoechst 33342 was used for counterstaining for 10 minutes. The fluorescently labeled tissue was placed cut-face down into incubation chambers and mounted with Mowiol mounting solution. The Leica SP8 confocal microscope was used to collect a series of x-y-z images. The images were typically 194×194 µm x-y size and were collected along the z-axis at 1 µm step size through 50 µm of bone marrow tissue. The 40x objective with a numerical zoom of 4 (pixel size: 0.142 µm) was used. The lung was imaged using an x63 objective with a 2.5 digital zoom (pixel size: 0.145 µm).

To perform 3D ECM analysis, bone marrow samples were fixed in 4% PFA and 0.4% GA for 1 hour, then embedded in 4% agarose and cut into 250 µm-thick sections using a vibratome. These sections were then incubated overnight at 4°C with primary antibodies targeting laminin, type IV collagen, fibronectin, and fibrinogen, followed by corresponding secondary antibodies and Hoechst 33342 counterstaining. Series of x-y-z images of typically 46.13*46.13 µm x-y size were collected along the z-axis at 1 µm step size through 15-35 µm of sinusoid-associated megakaryocytes, using the 63x objective with a numerical zoom of 4 (pixel size: 0.09) from a Leica SP8 confocal microscope. For quantitative spatial analysis of ECM around megakaryocytes, the observations have been made on z-stacks as described by Voisin et al. (Voisin et al., 2010).The fluorescence was delineated on the maximum z-stack projection of half a megakaryocyte using Image J software.

Image processing to perform quantitative analyses of fluorescence profiles and fiber length are explained in Suppl. Figure 7. The random length measurement method uses random sampling to provide unbiased data on laminin/collagen fibers in a 3D cage. Measurements included intervals between different branching points throughout the cage, including branch ends. Processing involved five steps: 1) acquiring 3D images, 2) projecting onto 2D planar sections, 3) selecting random intersection points for measurement, 4) measuring intervals using ImageJ software, and 5) repeating the process for a representative dataset. Briefly, a binary threshold mask (0 for the background and 1 for the ECM network, threshold to reduce signal noise) was generated from the channel showing labeled collagen or laminin. The mask was then applied to the channel showing megakaryocytes, resulting in an image that corresponded to only an active ECM signal that colocalized with megakaryocytes. The mask was then applied to the collagen/laminin channel to determine the amount of fluorescent ECM within the mask. This enabled the determination of the cell area and the elimination of the signal outside the mask.

Intravital imaging

Pf4cre and Itgb1-/-/Itgb3-/- mice underwent intravital imaging. To visualize megakaryocytes and sinusoids, an AF488-conjugated anti-GPIX antibody derivative and Texas Red dextran 70kDa were intravenously injected, respectively. The skull bone marrow was observed using two-photon microscopy, following the procedure described in reference (Bornert et al., 2021). The anesthetized mice were monitored for a maximum of 6 hours, during which one to four proplatelets were recorded. Two regions of interest were analyzed, each composed of 3 images (xyz = 1320 µm*1320 µm*100 µm, total volume/ROI = 0.17 µm3, pixel size = 0.867 µm at obj HC FLUOTAR L 25x/0.95 WATER zoom x1). A total of 130 megakaryocytes PF4cre vs. 130 megakaryocytes Itgb1-/-/Itgb3-/- were analyzed and the results are expressed in %.

Electron microscopy

For Transmission electron microscopy, bone marrow was fixed in 2.5% glutaraldehyde for 1 hour and embedded in Epon as described (Scandola et al., 2021). Transversal thin sections of the entire bone marrow were cut and examined under a JEOL TEM (120 kV). The number of megakaryocytes was counted per surface unit (s.u., 12,945 μm2). The observations from three independent mice were averaged.

For SEM, native bone marrow megakaryocytes were allowed to adhere to a surface coated with 300 µg/ml fibronectin, 100 µg/ml fibrinogen, or 50 µg/ml laminin 511. After gentle agitation to detach non-adherent cells, the remaining adherent cells were fixed in 2.5 % glutaraldehyde for 1 hour, dehydrated, attached to stubs, sputter coated, and examined under a Helios NanoLab microscope at 5kV (ThermoFisher, Eindhoven, The Netherlands). Adherent megakaryocytes where counted and classified as spreading or round cells. The results were obtained by the average from three independent mice.

Flow cytometry

Flow cytometry was used to investigate i) megakaryocyte ploidy, ii) intact megakaryocytes in the blood and iii) platelet activation.

For the ploidy analysis of the native megakaryocytes, mouse femurs, tibias, and iliac crests were harvested, cut into small fragments, and incubated in a PBS-collagenase-dispase mix (at 3 mg/mL and 4 mg/mL, 5mL/mouse respectively) for 30 min at 37°C. Then, the tube was filled with PBS-2% NCS to stop the collagenase, and the supernatant was 70 µm filtered to eliminate the bone fragments. After red cell lysis, the freshly isolated bone marrows were washed in PBS-2% NCS and stained with anti-CD41 and anti-CD42c antibodies and Hoechst. Anti Gr-1, B220, F4/80, and TER119 probes were used as a negative control to exclude granulocytes, B lymphocytes, macrophages, and erythrocytes from the total bone marrow suspension. The freshly isolated megakaryocytes were identified as a CD41/CD42c positive cell population and ploidy analysis was performed using Hoechst 33342. The results are representative of three independent experiments.

For detection of intact megakaryocytes in the blood, red blood cell lysis was performed using BD Lysing buffer solution according to manufacturer’s instruction. The resulting mononucleated cells were labelled with fluorescent conjugated antibodies: anti CD41-PE-Cy7 (clone MWReg30), anti-CD42c-Alexa 488 (clone RAM-1), anti-CD45-PE (clone 30-F11), TER119-APC (clone TER119), antibodies from ebiosciences except RAM-1 antibody produced in house. Megakaryocytes enriched from bone marrow using BSA gradient were spiked into blood sample to validate the gating strategy to identify potential circulating megakaryocytes as CD45+TER119-CD41+CD42c+ events. 300.103 CD45+TER119- events were recorded. Data acquisition was performed on a Symphony A1 flow cytometer (BD Biosciences) and analyzed using BD FACS Diva Software (BD Biosciences).

For platelet activation studies, whole blood was collected from the tail vein and anticoagulated with hirudin (200 U /ml). Platelets were activated or not with collagen-related peptide (CRP, 40 µg/ml) and TRAP (4mM) for 15 min at 37°C and stained with FITC-labelled or 647-labelled rat anti-mouse GPIbβ and APC-labeled rat anti-mouse P-selectin or PE-labelled rat anti-mouse activated-αIIbβ3 antibodies. Fluorescence was quantified using an LSRFortessa cell analyze (BD Biosciences) and BD FACSDiva software. A total of 20.000 events were analyzed for each sample.

Microfluidic experiments

The PDMS microfluidic chamber channels were assembled and connected to a peristaltic pump, as described (Osmani et al., 2021). The channels were washed with PBS for 5 min and coated with laminin 511 (50 µg/ml), fibrillar fibronectin (300 µg/ml), or fibrinogen (100 µg/ml) overnight at 4°C, followed by a blocking stage with human serum albumin (1 %) for 30 min at RT. To produce fibrillar fibronectin, we use the method of mechanical stretching as described in Maurer et al, (Maurer et al., 2015). Freshly isolated megakaryocytes were seeded at a concentration of 30 cells /µL per channel and incubated for 15 min – 45 min at 37 °C with 5 % CO2. In previous work, the seeding time was determined with capture efficiencies of over 80 % for Pf4cre megakaryocytes (Suppl. Figure 4B). The channels were then perfused under a flow of 300 µm/s and the megakaryocyte behavior was monitored in real-time, for 1 min, using a Leica DMI8 microscope (x20) equipped with a CMOS Camera (Hamamatsu ORCA fusion). Megakaryocyte capture yield was measured by quantifying the number of adherent megakaryocytes before and after flow induction (expressed as a percentage) on an average of 20 stage positions (n= 50 megakaryocytes analyzed/test). The results are an average of three independent experiments, the values were expressed as a mean ± sem.

Bone marrow explants

Preparation of bone marrow explants was performed as described in Guinard et al., 2021 to investigate proplatelet formation. For investigating megakaryocyte emigration, bone marrows were flushed from femurs of Pf4cre andItgb1-/-/Itgb3-/-; Lamα4+/+ and Lamα4-/-; as well as from DMSO- and B+I-treated mice and ten 0.5 mm-thick sections were placed in an incubation chamber. Megakaryocytes at the periphery of the explants were counted under an inverted phase contrast microscope coupled to a video camera (DMI8 Leica microscope, 40x objective). A motorized multiposition stage (in x, y, z) was used, and an average of 100 stages positions showing megakaryocytes was followed. Results are expressed as the number of total emigrating megakaryocytes at the indicated time points. In each case, a minimum of six independent experiments were performed. Mice from each genotype were analyzed and image acquisitions were performed in a blinded manner.

Effects of MMP inhibition

Batimastat plus ilomastat treatment was tested for their potential impact on the ECM in the microenvironment of the megakaryocyte vascular niche. Mice were treated daily with a protease inhibitor cocktail (Batimastat at 30 mg/kg + Ilomastat at 50 mg/kg) or vehicle (DMSO) for 7 consecutive days (Gui et al., 2018; Pielecka-Fortuna et al., 2015). The platelet count was monitored every two days. On the eighth day, the mice were sacrificed, and their bone marrow was collected for analysis of ECM organization and megakaryocyte behavior, as explained in the 3D confocal analysis on bone marrow vibratome sections.

Gelatin degradation assay

Coverslips were coated with Oregon green gelatin, and fixed with 0.5% glutaraldehyde for 20 min at RT. After washing three times with PBS, cells were seeded on coated coverslips and incubated for 6 h before fixation and staining.

Statistics

All values are reported as the mean ± sem. n= number of megakaryocytes studied. Statistical analyses were performed with PrismGraphpad software (La Jolla, CA, USA). For group comparison, data were tested for Gaussian distribution. Then, a Student t-test (Gaussian) or Mann-Whitney U test (non-Gaussian) were used to compare individual groups; multiple groups were compared by one-way ANOVA followed by Bonferroni post-test or by a non-parametric Kruskal-Wallis test, with a threshold of significance of 5%. P-values <0.05 were considered statistically significant. *P<0.05; **P<0.01; ***P<0.001.